```python

```

<img align="center" src="images/course.png" width="800">

# 16720 (B)  3D Reconstruction - Assignment 5 - q3
    Instructor: Kris                          TAs: Zen (Lead), Yan, Rawal, Wen-Hsuan, Paritosh, Qichen

```python
# Helper functions for this assignment. DO NOT MODIFY!!!
"""
Helper functions.

Written by Chen Kong, 2018.
Modified by Zhengyi (Zen) Luo, 2021
"""

import numpy as np
import matplotlib.pyplot as plt
import matplotlib
import scipy
import cv2
import nbimporter
from q1 import eightpoint, sevenpoint, camera2, _epipoles, calc_epi_error, toHomogenous, _singularize
from q2 import find_M2, plot_3D

def plot_3D_dual(P_before, P_after):
    matplotlib.use('TkAgg')
    fig = plt.figure()
    ax = fig.add_subplot(111, projection='3d')
    ax.set_title("Blue: before; red: after")
    ax.scatter(P_before[:,0], P_before[:,1], P_before[:,2], c = 'blue')
    ax.scatter(P_after[:,0], P_after[:,1], P_after[:,2], c='red')
    while True:
        x, y = plt.ginput(1, mouse_stop=2)[0]
        plt.draw()


```

## Q3: Bundle Adjustment
Bundle Adjustment is commonly used as the last step of every feature-based 3D reconstruction algorithm. Given a set of images depicting a number of 3D points from different viewpoints, bundle adjustment is the process of simultaneously refining the 3D coordinates along with the camera parameters. It minimizes reprojection error, which is the squared sum of distances between image points and predicted points. In this section, you will implement bundle adjustment algorithm by yourself. Specifically,


- In Q3.1, you need to implement a RANSAC algorithm to estimate the fundamental matrix F and all the inliers.
- In Q3.2, you will need to write code to parameterize Rotation matrix $\mathbf{R}$ using [Rodrigues formula](https://en.wikipedia.org/wiki/Rodrigues\%27\_formul) (Please check [this pdf](https://www2.cs.duke.edu/courses/fall13/compsci527/notes/rodrigues.pdf) for a detailed explanation), which will enable the joint optimization process for Bundle Adjustment.
- Q3.3, you will need to first write down the objective function in rodriguesResidual, and do the bundleAdjustment.

### Q3.1 RANSAC for Fundamental Matrix Recovery (15 pt implementation)

In some real world applications, manually determining correspondences is infeasible and often there will be noisy correspondences. Fortunately, the RANSAC method seen (and implemented in previous assignments) in class can be applied to the problem of fundamental matrix estimation.

Implement the above algorithm with the signature:
```
[F, inliers] = ransacF(pts1, pts2, M)
```

where `M` is defined in the same way as when we calculate the fundamental matrix and inliers is a boolean vector of size equivalent to the number of points. Here inliers are set to true only for the points that satisfy the threshold defined for the given fundamental matrix F.

We have provided some noisy coorespondances in some\_corresp\_noisy.npz in which around $75\%$ of the points are inliers. Compare the result of RANSAC with the result of the eight-point algorithm when ran on the noisy correspondences. 

**Hints:** Use the seven point to compute the fundamental matrix from the minimal set of points. Then compute the inliers, and refine your estimate using all the inliers.


<!-- grade_id: cell-450fffa620eaf4bc, points: 0, grade: false, solution: true -->
```python
def ransacF(pts1, pts2, M):
    '''
    Q3.1: RANSAC method.
        Input:  pts1, Nx2 Matrix
                pts2, Nx2 Matrix
                M, a scalar parameter
        Output: F, the fundamental matrix
                inlier_curr, Nx1 bool vector set to true for inliers
    ***
    Hints:
    (1) You can use the calc_epi_error from q1 with threshold to calculate inliers. Tune the threshold based on 
        the results/expected number of inliers. You can also define your own metric. 
    (2) Use the seven point algorithm to estimate the fundamental matrix as done in q1
    (3) Choose the resulting F that has the most number of inliers
    (4) You can increase the nIters to bigger/smaller values
        
    '''
    N = pts1.shape[0]
    pts1_homo, pts2_homo = toHomogenous(pts1), toHomogenous(pts2)
    max_iterations = 500
    best_inlier = 0
    F = None
    best_inlier_mask = None

    # ----- TODO -----
    # YOUR CODE HERE

    NUM_REQUIRED_INLIERS = 10
    # Initial threshold is for checking the sevenpoint pass
    INITIAL_THRESHOLD = 16
    # Tighter threshold is for checking the eightpoint pass with more points
    TIGHTER_THRESHOLD = 3
    # If we hit this fraction of inliers, just end it
    CONVERGE = 0.8

    indices = np.array(range(pts1.shape[0]))
    for i in range(max_iterations):
        if i % 100 == 0:
            print(f"Iteration {i}")
        index_option = np.random.choice(indices, size=7, replace=False)
        F_array_option = sevenpoint(pts1[index_option], pts2[index_option], M)
        for F_option in F_array_option:
            error = calc_epi_error(pts1_homo, pts2_homo, F_option)
            inlier_mask_option = error < INITIAL_THRESHOLD
            if sum(inlier_mask_option) < NUM_REQUIRED_INLIERS:
                continue
            F_better = eightpoint(pts1[inlier_mask_option], pts2[inlier_mask_option], M)
            error_better = calc_epi_error(pts1_homo, pts2_homo, F_better)
            inlier_mask_better = error_better < TIGHTER_THRESHOLD
            num_inliers = sum(inlier_mask_better)
            if num_inliers > best_inlier:
                print(f"Found model with {num_inliers}/{len(pts1)} inliers, better than {best_inlier}."
                      f" Average inlier error {np.average(error_better[inlier_mask_better])} px^2")
                F = F_better
                best_inlier = num_inliers
                best_inlier_mask = inlier_mask_better
            # Just return if we get sufficient inliers
            if num_inliers / len(pts1) > CONVERGE:
                print(f"Returning due to high inlier count")
                return F, best_inlier_mask

    return F, best_inlier_mask


np.random.seed(0)
np.set_printoptions(precision=4, suppress=1)
correspondence = np.load('data/some_corresp_noisy.npz') # Loading noisy correspondences
intrinsics = np.load('data/intrinsics.npz') # Loading the intrinscis of the camera
K1, K2 = intrinsics['K1'], intrinsics['K2']
pts1, pts2 = correspondence['pts1'], correspondence['pts2']
im1 = plt.imread('data/im1.png')
im2 = plt.imread('data/im2.png')
F, inliners = ransacF(pts1, pts2, M=np.max([*im1.shape, *im2.shape]))

```

Output:
```
Iteration 0
Found model with 97/140 inliers, better than 0. Average inlier error 0.629273329591031 px^2
Found model with 103/140 inliers, better than 97. Average inlier error 0.4779824277678971 px^2
Iteration 100
Found model with 105/140 inliers, better than 103. Average inlier error 0.4706820710493711 px^2
Iteration 200
Found model with 109/140 inliers, better than 105. Average inlier error 0.35434561915729035 px^2
Iteration 300
Iteration 400
```

<!-- grade_id: q3_1_a, points: 15, grade: true, solution: false -->
```python
# Full set of tests; you will get full points for coding if passing the following tests. 
assert(np.sum(inliners) > len(pts1) * 0.7)
```

### Q3.2 Rodrigues and Invsere Rodrigues (10 pt implementation)
So far we have independently solved for the camera matrix, $\mathbf{M}_j$ and 3D projections, $\textbf{w}_i$. In bundle adjustment, we will jointly optimize the reprojection error with respect to the points $\textbf{w}_i$ and the camera matrix $\textbf{w}_j$.

$$
err = \sum_{ij} ||\textbf{x}_{ij} - Proj(\mathbf{C}_j, \textbf{w}_i)||^2,
$$
where $\textbf{w}_j = \mathbf{K}_j \mathbf{M}_j$.

For this homework, we are going to only look at optimizing the extrinsic matrix. The rotation matrix forms the Lie Group $\textbf{SO}(3)$ that doesn't satisfy the addition operation so it cannot be directly optimized. Instead, we parameterize the rotation matrix to axis angle using Rodrigues formula to the Lie Algebra $\mathfrak{so}(3)$, which is defined in $\mathbb{R}^3$. through which the least squares optimization process can be done to optimize the axis angle. Try to implement function

```
R = rodrigues(r)
```

as well as the inverse function that converts a rotation matrix $\mathbf{R}$ to a Rodrigues vector $\mathbf{r}$

```
r = invRodrigues(R)
```

Please refer to [Rodrigues formula](https://en.wikipedia.org/wiki/Rodrigues\%27\_formul)  and [this pdf](https://www2.cs.duke.edu/courses/fall13/compsci527/notes/rodrigues.pdf) for reference.


<!-- grade_id: cell-f7107c598b84fe13, points: 0, grade: true, solution: true -->
```python
def rodrigues(r):
    '''
    Q3.2: Rodrigues formula.
        Input:  r, a 3x1 vector
        Output: R, a 3x3 rotation matrix
    '''
    # ----- TODO -----
    # YOUR CODE HERE
    if r.shape == (3,):
        r = r.reshape((3, 1))
    assert r.shape == (3, 1)

    theta = np.linalg.norm(r)
    unit = r / theta
    # Define the matrix version of a cross-product
    skew = np.array([
        [ 0,          -unit[2, 0],  unit[1, 0]],
        [ unit[2, 0],  0,          -unit[0, 0]],
        [-unit[1, 0],  unit[0, 0],  0],
    ])
    
    R = np.eye(3) * np.cos(theta) + \
        (1 - np.cos(theta)) * unit @ unit.T + \
        skew * np.sin(theta)
    return R


def invRodrigues(R):
    '''
    Q5.2: Inverse Rodrigues formula.
        Input:  R, a 3x3 rotation matrix
        Output: r, a 3x1 vector
            WRONG - they want a (3,), it fails otherwise
    '''
    # ----- TODO -----
    # YOUR CODE HERE
    
    # Used implementation here instead of the PDF, it seems like there are some
    # missing conditions in the PDF
    # https://math.stackexchange.com/questions/83874/efficient-and-accurate-numerical-implementation-of-the-inverse-rodrigues-rotatio
    r = np.array([
        [R[2, 1] - R[1, 2]],
        [R[0, 2] - R[2, 0]],
        [R[1, 0] - R[0, 1]],
    ])
    t = np.trace(R)
    theta = np.arccos((t - 1) / 2)
    e = 1e-6
    
    # Create the sometimes necessary v vector
    a = np.argmax(np.diag(R))
    b = (a + 1) % 3
    c = (a + 2) % 3
    v = np.zeros((3, 1))
    s = np.sqrt(R[a, a] - R[b, b] - R[c, c] + 1)
    v[a] = s / 2
    v[b] = (1 / (2 * s)) * (R[b, a] + R[a, b])
    v[c] = (1 / (2 * s)) * (R[c, a] + R[a, c])

    if t >= (3 - e):
        w = ((1/2) - (t-3)/12) * r
    elif t > (-1 + e):
        w = (theta / (2 * np.sin(theta))) * r
    else:
        w = np.pi * v / np.linalg.norm(v)
    return w.reshape(3)


# # Check for orthonormality
# R = rodrigues(np.array([[0.4], [2.5], [-1.1]]))
# for i in range(3):
#     assert np.isclose(np.linalg.norm(R[:, i]), 1.0)
# assert np.isclose(R[:, 0] @ R[:, 1], 0.0)
# assert np.isclose(R[:, 0] @ R[:, 2], 0.0)
# assert np.isclose(R[:, 1] @ R[:, 2], 0.0)

# # Try to extract the old values
# r = invRodrigues(R)
# print(f"Extracted values {r.T}")
```

<!-- grade_id: q3_2_a, points: 3, grade: true, solution: false -->
```python
# Simple Tests to verify your implmentation:
from scipy.spatial.transform import Rotation as sRot
rotVec = sRot.random()
mat = rodrigues(rotVec.as_rotvec())

assert(np.linalg.norm(rotVec.as_rotvec() - invRodrigues(mat)) < 1e-3)
assert(np.linalg.norm(rotVec.as_matrix() - mat) < 1e-3)


```

<!-- grade_id: q3_2_b, points: 7, grade: true, solution: false -->
```python
# Hidden Tests
```

### Q3.3 Bundle Adjustment (10 pt writeup)

In this section, you need to implement the bundle adjustment algorithm. Using the parameterization you implemented in the last question, write an objective function for the extrinsic optimization:

```
residuals = rodriguesResidual(K1, M1, p1, K2, p2, x)
```
where x is the flattened concatenation of $\mathbf{w}$, $\mathbf{r}_2$, and $\mathbf{t}_2$.
$\mathbf{w}$ are the 3D points; $\mathbf{r}_2$ and $\mathbf{t}_2$ are the rotation (in the Rodrigues vector form) and translation vectors associated with the projection matrix $\mathbf{M}_2$; $p1$ and $p2$ are 2D coordinates of points in image 1 and 2, respectively. The `residuals` are the difference between the original image projections and the estimated projections (the square of $2$-norm of this vector corresponds to the error we computed in Q3.2):
```
residuals = numpy.concatenate([(p1-p1').reshape([-1]), (p2-p2').reshape([-1])])
```

Use this objective function and Scipy's nonlinear least squares optimizer $\texttt{leastsq}$ write a function to optimize for the best extrinsic matrix and 3D points using the inlier correspondences from some_corresp_noisy.npz and the RANSAC estimate of the extrinsics and 3D points as an initialization.

```
[M2, w, o1, o2] = bundleAdjustment(K1, M1, p1, K2, M2_init, p2, p_init)
```

Try to extract the rotation and translation from `M2_init`, then use `invRodrigues` you implemented previously to transform the rotation, concatenate it with translation and the 3D points, then the concatenate vector are variables to be optimized. After obtaining optimized vector, decompose it back to rotation using `Rodrigues` you implemented previously, translation and 3D points coordinates.

<span style='color:red'>**Output:**</span> In your write-up: include an image of output of the `plot_3D_dual` function by passing in the original 3D points and the optimized points. Also include the before and after reprojection error for the `rodriguesResidual` function.



<!-- grade_id: cell-21b599d494aeefec, points: 0, grade: false, solution: true -->
```python
def extract_flat_M2(x_vector):
    r = x_vector[-6:-3].reshape((3, 1))
    R = rodrigues(r)
    T = x_vector[-3:].reshape((3, 1))
    return np.hstack((R, T))


def rodriguesResidual(x, K1, M1, p1, K2, p2):
    '''
    Q3.3: Rodrigues residual.
        Input:  x, the flattened concatenation of P, r2, and t2.
                K1, the intrinsics of camera 1
                M1, the extrinsics of camera 1
                p1, the 2D coordinates of points in image 1
                K2, the intrinsics of camera 2
                p2, the 2D coordinates of points in image 2
        Output: residuals, 4N x 1 vector, the difference between original 
                and estimated projections
    '''
    residuals = None
    # ----- TODO -----
    # YOUR CODE HERE

    # Recover the 3D points
    points3d = x[:-6].reshape((-1, 3))
    p3d_homogeneous = np.hstack([points3d,
                                 np.ones((points3d.shape[0], 1))])
    M2 = extract_flat_M2(x)
    
    p1_projogeneous = (K1 @ M1 @ p3d_homogeneous.T).T
    p2_projogeneous = (K2 @ M2 @ p3d_homogeneous.T).T

    p1_proj = (p1_projogeneous[:, :2] / p1_projogeneous[:, 2:])
    p2_proj = (p2_projogeneous[:, :2] / p2_projogeneous[:, 2:])
    
    residuals = np.concatenate([(p1 - p1_proj).reshape([-1]),
                                (p2 - p2_proj).reshape([-1])])
    return residuals


def bundleAdjustment(K1, M1, p1, K2, M2_init, p2, P_init):
    '''
    Q3.3 Bundle adjustment.
        Input:  K1, the intrinsics of camera 1
                M1, the extrinsics of camera 1
                p1, the 2D coordinates of points in image 1
                K2,  the intrinsics of camera 2
                M2_init, the initial extrinsics of camera 1
                p2, the 2D coordinates of points in image 2
                P_init, the initial 3D coordinates of points
        Output: M2, the optimized extrinsics of camera 1
                P2, the optimized 3D coordinates of points
                o1, the starting objective function value with the initial input
                o2, the ending objective function value after bundle adjustment
    
    ***
    Hints:
    (1) Use the scipy.optimize.minimize function to minimize the objective function, rodriguesResidual. 
        You can try different (method='..') in scipy.optimize.minimize for best results. 
    '''
    obj_start = obj_end = 0
    # ----- TODO -----
    # YOUR CODE HERE
    
    # Build our Frankenstein x-vector
    x0 = np.concatenate([P_init.flatten(),
                         invRodrigues(M2_init[:3, :3]),
                         M2_init[:3,  3].flatten()])

    args = (K1, M1, p1, K2, p2)
    kwargs = {"K1": K1, "M1": M1, "p1": p1, "K2": K2, "p2": p2}
    optimized = scipy.optimize.minimize(
        fun=lambda *x: sum((rodriguesResidual(*x))**2),
        x0=x0,
        args=args,
    )
    # minimize appeared to return better results, least_squares results were
    # very scaled. Is there an arbitrary scale factor involved?
    # optimized = scipy.optimize.least_squares(
    #     fun=rodriguesResidual,
    #     x0=x0,
    #     kwargs=kwargs,
    # )

    M2 = extract_flat_M2(optimized.x)
    P = optimized.x[:-6].reshape((-1, 3))
    assert P.shape == P_init.shape

    obj_start = sum(rodriguesResidual(x0, **kwargs)**2)
    obj_end = sum(rodriguesResidual(optimized.x, **kwargs)**2)
    
    return M2, P, obj_start, obj_end
```

```python
# Visualization:
np.random.seed(0)
correspondence = np.load('data/some_corresp_noisy.npz') # Loading noisy correspondences
intrinsics = np.load('data/intrinsics.npz') # Loading the intrinscis of the camera
K1, K2 = intrinsics['K1'], intrinsics['K2']
pts1, pts2 = correspondence['pts1'], correspondence['pts2']
im1 = plt.imread('data/im1.png')
im2 = plt.imread('data/im2.png')

M=np.max([*im1.shape, *im2.shape])
F, inliners = ransacF(pts1, pts2, M)
pts1_inliners = pts1[inliners.squeeze(), :]
pts2_inliners = pts2[inliners.squeeze(), :]

M1 = np.hstack((np.identity(3), np.zeros(3)[:,np.newaxis]))
C1 = K1.dot(M1)
M2_init,C2, P_init = find_M2(F, pts1_inliners, pts2_inliners, intrinsics)

M2, P_final, obj_start, obj_end = bundleAdjustment(K1, M1, pts1_inliners, K2, M2_init, pts2_inliners, P_init)
print(f"Before {obj_start}, After {obj_end}")
plot_3D_dual(P_init, P_final)

```

Output:
```
Iteration 0
Found model with 97/140 inliers, better than 0. Average inlier error 0.629273329591031 px^2
Found model with 103/140 inliers, better than 97. Average inlier error 0.4779824277678971 px^2
Iteration 100
Found model with 105/140 inliers, better than 103. Average inlier error 0.4706820710493711 px^2
Iteration 200
Found model with 109/140 inliers, better than 105. Average inlier error 0.35434561915729035 px^2
Iteration 300
Iteration 400
Best Error 347.1649703432302
Before 347.1649703432219, After 9.515508068588574
~/.local/lib/python3.8/site-packages/matplotlib/backends/_backend_tk.py in blit(photoimage, aggimage, offsets, bbox)
     73         bboxptr = (x1, x2, y1, y2)
     74     else:
---> 75         photoimage.blank()
     76         bboxptr = (0, width, 0, height)
     77     _tkagg.blit(

/usr/lib/python3.8/tkinter/__init__.py in blank(self)
   4066     def blank(self):
   4067         """Display a transparent image."""
-> 4068         self.tk.call(self.name, 'blank')
   
```

```python

```